# TD - Fonctions logiques (suite)
### 3. Ecrire une fonction logique à partir de son expression.
Considérons la fonction logique décrite par l'expression ci dessous :
$$S = (\overline{A} + B) \times C $$

Ecrire une fonction `logique1(A,B,C)` qui renvoie le résultat de cette expression.

<!-- grade_id: cell-137d01c21c7c4f91, points: 0, grade: false, solution: true -->
```python
def logique1(A,B,C) :
    ### votre code ici
    s=
    return s
```

Vérifions maintenant que votre fonction passe bien le jeu de tests :

<!-- grade_id: cell-cc6fd57a98b27ab4, points: 2, grade: true, solution: false -->
```python
assert logique1(False,False,False) == False
assert logique1(False,True,True) == True
assert logique1(False,False,True) == True
assert logique1(True,True,False) == False
assert logique1(True,True,True)==True
assert logique1(True,False,True)==False
assert logique1(True,True,False)==False
```

### 4. Ecrire une fonction logique à partir de sa table de vérité.
Considérons la table de vérité ci dessous :

|  A  |  B  |  C  |  S  |
|:---:|:---:|:---:|:---:|
|0|0|0|1|
|0|0|1|0|
|0|1|0|1|
|0|1|1|0|
|1|0|0|1|
|1|0|1|0|
|1|1|0|1|
|1|1|1|1|

- Trouver l'expression logique en fonction de A,B,C qui vérifie cette table de vérité

S = 

- Puis écrire une fonction `logique2(A,B,C)` qui renvoie une valeur correspondant à cette expression logique.

<!-- grade_id: cell-1cd5540bfe60d066, points: 0, grade: false, solution: true -->
```python
def logique2(A,B,C) :
    ### votre code ici
    x= 
    return x
```

- Vérifier enfin que votre fonction passe bien la batterie de tests suivante :

<!-- grade_id: cell-2782b2aa2a543221, points: 3, grade: true, solution: false -->
```python
assert logique2(False,False,False) == True
assert logique2(False,True,True) == False
assert logique2(True,True,False) == True
assert logique2(True,True,True) == True
assert logique2(False,False,True)==False
assert logique2(False,True,False)==True
assert logique2(True,False,False)==True
assert logique2(True,False,True)==False
```

### 5. Ecrire une fonction logique à partir d'un circuit logique
Considérons le circuit ci-dessous
![20% center](circuit1.PNG)
Ecrire une fonction `logique3(A,B,C)` qui renvoie une valeur qui corresponde à ce circuit.

<!-- grade_id: cell-d726d50776cc2830, points: 0, grade: false, solution: true -->
```python
def logique3(A,B,C) :
    ### votre code ici
    x= 
    return x
```

Quelques tests qui vont vérifier la conformité de votre fonction.

<!-- grade_id: cell-c5a80c7c71ebac29, points: 2, grade: true, solution: false -->
```python
assert logique3(False,False,False) == True
assert logique3(False,True,True) == True
assert logique3(True,False,True) == True
assert logique3(True,True,True) == False
assert logique3(False,False,True) == True
assert logique3(False,True,False) == True
assert logique3(True,False,False) == True
assert logique3(True,True,False) == True
```